# Assignment 1
For this assignment you are welcomed to use other regex resources such a regex "cheat sheets" you find on the web. Feel free to share good resources with your peers in slack!



Before start working on the problems, here is a small example to help you understand how to write your own answers. In short, the solution should be written within the function body given, and the final result should be returned. Then the autograder will try to call the function and validate your returned result accordingly. 

```python
def example_word_count():
    # This example question requires counting words in the example_string below.
    example_string = "Amy is 5 years old"
    
    # YOUR CODE HERE.
    # You should write your solution here, and return your result, you can comment out or delete the
    # NotImplementedError below.
    result = example_string.split(" ")
    return len(result)

    #raise NotImplementedError()
```

## Part A

Find a list of all of the names in the following string using regex.

<!-- grade_id: names, points: 0, grade: false, solution: true -->
```python
import re
def names():
    simple_string = """Amy is 5 years old, and her sister Mary is 2 years old. 
    Ruth and Peter, their parents, have 3 kids."""

    ### BEGIN SOLUTION  
    pattern = r'[A-Z][a-z]*'
    match = re.findall(pattern, simple_string)
    return match

    # Alternative answers: 
    # pattern = r'[A-Z]\w+'
    # pattern = r'[A-Z]\S+'
    
    # Hints:
    # hint1: names start with capital letters
    # hint2: how do you generate a list of matched chunks? :
    # use re.findall(pattern, string) 
    # hint3: ????? <THERE SEEMS TO BE NO HINT 3?>
    # hint4: how do you find one or more of particular tokens?
    # use +
    
    ### END SOLUTION
```

<!-- grade_id: correct_names, points: 1, grade: true, solution: false -->
```python
assert len(names()) == 4, "There are four names in the simple_string"

### BEGIN HIDDEN TESTS
assert ((len(names()[0])>2)&(len(names()[0])>2)&(len(names()[0])>2)&(len(names()[0])>2)) == True, "Please match all the characters in these names."

assert ('Amy' in names())== True
assert ('Mary' in names())== True
assert ('Ruth' in names())== True
assert ('Peter' in names())== True
### END HIDDEN TESTS
```

## Part B

The dataset file in [assets/grades.txt](assets/grades.txt) contains a line separated list of people with their grade in 
a class. Create a regex to generate a list of just those students who received a B in the course.

<!-- grade_id: grades, points: 0, grade: false, solution: true -->
```python
import re
def grades():
    with open ("assets/grades.txt", "r") as file:
        grades = file.read()

    ### BEGIN SOLUTION
    pattern = re.compile(r'\w+\s\w+(?=: B)')
    matches = re.findall(pattern,grades)
    return matches  

    # Alternative answers: 
    # pattern = """(?P<test>\w+\s+\w+): B"""
    
    # Hints:
    # hint1: What is look ahead in RegEx and how do you use it?
    # if you want to find A in front of B, the way how you write a look ahead regex is as follows:
    # r'A(?=B)'
    # hint2: how do you generate a list of matched chunks? :
    # use re.findall(pattern, string) 
    # hint3: how do you include both spaces and words?
    # use both \w and \s
    # hint4: how do you find one or more of particular tokens?
    # use +    

    ### END SOLUTION
```

<!-- grade_id: correct_grades, points: 1, grade: true, solution: false -->
```python
assert len(grades()) == 16

### BEGIN HIDDEN TESTS

name_list = ['Bell Kassulke',
 'Simon Loidl',
 'Elias Jovanovic',
 'Hakim Botros',
 'Emilie Lorentsen',
 'Jake Wood',
 'Fatemeh Akhtar',
 'Kim Weston',
 'Yasmin Dar',
 'Viswamitra Upandhye',
 'Killian Kaufman',
 'Elwood Page',
 'Elodie Booker',
 'Adnan Chen',
 'Hank Spinka',
 'Hannah Bayer']

assert (': B' in grades()[1]) == False, "Please only return student names."
assert ('Kassulke' in grades()) == False, "Please return students' full name."
for name in name_list:
    assert (name in grades())== True, "Check your name list."
    
### END HIDDEN TESTS
```

## Part C

Consider the standard web log file in [assets/logdata.txt](assets/logdata.txt). This file records the access a user makes when visiting a web page (like this one!). Each line of the log has the following items:
* a host (e.g., '146.204.224.152') 
* a user_name (e.g., 'feest6811' **note: sometimes the user name is missing! In this case, use '-' as the value for the username.**)
* the time a request was made (e.g., '21/Jun/2019:15:45:24 -0700')
* the post request type (e.g., 'POST /incentivize HTTP/1.1' **note: not everything is a POST!**)

Your task is to convert this into a list of dictionaries, where each dictionary looks like the following:
```
example_dict = {"host":"146.204.224.152", 
                "user_name":"feest6811", 
                "time":"21/Jun/2019:15:45:24 -0700",
                "request":"POST /incentivize HTTP/1.1"}
```

<!-- grade_id: logs, points: 0, grade: false, solution: true -->
```python
import re
def logs():
    with open("assets/logdata.txt", "r") as file:
        logdata = file.read()
    
    ### BEGIN SOLUTION
    pattern = """
        (?P<host>[\w\.]+)             # host %h
        \s\-\s(?P<user_name>.*)\s     # user %u
        \[(?P<time>.+)\]              # time %t
        \s\"(?P<request>.*)\"         # request "%r"
    """
    result = [item.groupdict() for item in re.finditer(pattern,logdata,re.VERBOSE)]
    
    return result
    
    ## Hint:

    ## hint1: What is look ahead in RegEx and how do you use it?
    ## if you want to find A in front of B, the way how you write a look ahead regex is as follows:
    ## r'A(?=B)'

    ## hint2: how do you generate a list of matched chunks? :
    ## use re.findall(pattern, string) 

    ## hint3: how do you include both spaces and words?
    ## use both \w and \s

    ## hint4: how do you find one or more of particular tokens?
    ## use +

    ## hint5: how do you ignore paricular tokens using RegEx look-ahead?
    ## Suppose you have string 'AB' don't want to include 'A' and want to detect 'B' only. In this case, you can either
    ## A(?P<test>B) or 
    ## A
    ## (?P<test>B)

    ## hint6: how do you use write a long regex with comments and linebreaks?
    ## use re.VERBOSE    
    
    ### END SOLUTION
```

<!-- grade_id: cell-correct_logs, points: 1, grade: true, solution: false -->
```python
assert len(logs()) == 979

one_item={'host': '146.204.224.152',
  'user_name': 'feest6811',
  'time': '21/Jun/2019:15:45:24 -0700',
  'request': 'POST /incentivize HTTP/1.1'}
assert one_item in logs(), "Sorry, this item should be in the log results, check your formating"

### BEGIN HIDDEN TESTS
assert len(logs()[0:1][0])==4, "Make sure that each of your dictionary is in the right format. "

tmp3 = [{'host': '146.204.224.152',
  'user_name': 'feest6811',
  'time': '21/Jun/2019:15:45:24 -0700',
  'request': 'POST /incentivize HTTP/1.1'},
 {'host': '98.143.107.61',
  'user_name': 'reilly5813',
  'time': '21/Jun/2019:15:54:04 -0700',
  'request': 'HEAD /channels HTTP/1.1'},
 {'host': '93.88.34.132',
  'user_name': '-',
  'time': '21/Jun/2019:15:49:04 -0700',
  'request': 'POST /24%2f365 HTTP/2.0'},
 {'host': '100.32.205.59',
  'user_name': 'ortiz8891',
  'time': '21/Jun/2019:15:45:28 -0700',
  'request': 'PATCH /architectures HTTP/1.0'},
 {'host': '168.95.156.240',
  'user_name': 'stark2413',
  'time': '21/Jun/2019:15:45:31 -0700',
  'request': 'GET /engage HTTP/2.0'},
 {'host': '72.157.89.10',
  'user_name': '-',
  'time': '21/Jun/2019:15:47:14 -0700',
  'request': 'PATCH /clicks-and-mortar/metrics HTTP/1.0'},
 {'host': '180.95.121.94',
  'user_name': 'mohr6893',
  'time': '21/Jun/2019:15:45:34 -0700',
  'request': 'PATCH /extensible/reinvent HTTP/1.1'},
 {'host': '144.23.247.108',
  'user_name': 'auer7552',
  'time': '21/Jun/2019:15:45:35 -0700',
  'request': 'POST /extensible/infrastructures/one-to-one/enterprise HTTP/1.1'},
 {'host': '180.234.185.134',
  'user_name': 'windler5463',
  'time': '21/Jun/2019:15:46:34 -0700',
  'request': 'POST /synergies/innovative/enterprise HTTP/1.1'},
 {'host': '233.187.15.207',
  'user_name': '-',
  'time': '21/Jun/2019:15:46:14 -0700',
  'request': 'GET /harness/intuitive HTTP/1.0'}]

for dic in tmp3:
    assert dic in logs(), "Please include all the records"
### END HIDDEN TESTS
```